Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Athish Prakalath V P"
```

# Agent Based Modeling - Mesa

"[Mesa](https://mesa.readthedocs.io/en/stable/index.html) is an agent-based modeling framework in Python. Mesa allows users to quickly create agent-based models using built-in core components (such as spatial grids and agent schedulers) or customized implementations, visualize them using a browser-based interface, and analyze their results using Python’s data analysis tools."

_This notebook is based on and uses parts of texts from the Mesa docs._

### Discretised Lotka Volterra
In this notebook we aim to familiarise you with Mesa and its Object Oriented Programming approach to agent-based models. This will be done through implementation of a simplified preditor-prey model (discretised Lotka Volterra). 

If you are not yet comfortable with Object Oriented Programming or need a refresher, please visit [this page](https://python.swaroopch.com/oop.html).

We will create a version of the "sheep-wolf" model wherein wolves and sheep wander randomly around the landscape. To allow the population to grow, each wolf or sheep has a fixed probability of reproducing at each timestep. In this simplified version, sheep do not have to eat, and only die when eaten by a wolf. Wolves, however, have a probability of dying. This produces interesting population dynamics, but creates a model that is ultimately unstable.

To do this we need the following:
- **Agents:** the predators and preys
- **Environment:** an area in which the predators and prey can move around and interact
- **Model:** the controller that links the environment and the agents

We've already implemented the model and the environment is already implemented by Mesa, but it is up to you to implement the Agents!

```python
from IPython.display import clear_output
import mesa
#import ipynb
clear_output(wait=True)
print("Everything imported!")
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[6], line 2
      1 from IPython.display import clear_output
----> 2 import mesa
      3 #import ipynb
      4 clear_output(wait=True)

ModuleNotFoundError: No module named 'mesa'
```

### Testing
Since testing each of the methods that you will create seperately might be difficult, we provide you with a very minimal model that will be used during the testing phase. Later in this notebook we will introduce you to the model that has been implemented in `model.py`.

```python
from mesa import Model
from mesa.space import MultiGrid

class TestModel(Model):
    def __init__(self, width, height):
        self.height = height
        self.width = width
        
        self.grid = MultiGrid(self.width, self.height, torus=True)
        
        self.n_agents = 0
        self.agents = []

    
    def new_agent(self, agent_type, pos):
        '''
        Method that enables us to add agents of a given type.
        '''
        self.n_agents += 1
        
        # Create a new agent of the given type
        new_agent = agent_type(self.n_agents, self, pos)
        
        # Place the agent on the grid
        self.grid.place_agent(new_agent, pos)
        
        # And add the agent to the model so we can track it
        self.agents.append(new_agent)
        
    def remove_agent(self, agent):
        '''
        Method that enables us to remove passed agents.
        '''
        self.n_agents -= 1
        
        # Remove agent from grid
        self.grid.remove_agent(agent)
        
        # Remove agent from model
        self.agents.remove(agent)
        
    def step(self):
        '''
        Method that steps every agent. 
        
        Prevents applying step on new agents by creating a local list.
        '''
        for agent in list(self.agents):
            agent.step()

```

### Agents - Random walker
As both the sheep and the wolves will walk around randomly, we will first implement a `RandomWalker` class. We can then re-use (inherit) this class when making our `Wolf` and `Sheep` classes. Below you can find the framework for the class, but you will have to implement the `random_move()` method. This method should get the neighbouring cells (Moore's neighbourhood), select one, and move the agent to this cell.

Find the following functions through the provided link and find out which inputs they need:
  - [`MultiGrid.get_neighborhood()`](https://mesa.readthedocs.io/en/stable/apis/space.html#mesa.space.MultiGrid.get_neighborhood)
 - [`MultiGrid.move_agent()`](https://mesa.readthedocs.io/en/stable/apis/space.html#mesa.space.MultiGrid.move_agent)
 
Keep in mind that the `model` that is passed when the RandomWalker is initialised contains the grid. It is accessable through `self.model.grid` in the `random_move` function. 

<!-- grade_id: cell-1bff8a60c7a9567e, points: 0, grade: false, solution: true -->
```python
from mesa import Agent
import random

class RandomWalker(Agent):
    def __init__(self, unique_id, model, pos):
        super().__init__(unique_id, model)

        self.pos = pos

    def random_move(self):
        ''' 
        This method should get the neighbouring cells (Moore's neighbourhood), select one, and move the agent to this cell.
        '''
        possible_steps = self.model.grid.get_neighborhood(
            self.pos,
            moore=True,
            include_center=False,
        )
        new_position = random.choice(possible_steps)
        self.model.grid.move_agent(self, new_position)

```

<!-- grade_id: cell-8b1f4b5f5975c03b, points: 1, grade: true, solution: false -->
```python
tester = TestModel(10, 10)

# Create a RandomWalker, so that we can call the random_move() method
start_position = (5, 5)
tester.new_agent(RandomWalker, start_position)

# Create a reference, so that we can properly test
walker = tester.agents[0]

# Check movement
walker.random_move()
assert walker.pos != (5, 5), "The walker has not moved!"
assert abs(walker.pos[0] - 5) <= 1 and abs(walker.pos[1] - 5) <= 1, "The walker has moved more than 1 square"
```

### Agents - Sheep
As described above, to allow the population to grow, each wolf or sheep has a fixed probability of reproducing at each timestep. Sheep do not have to eat, and only die when eaten by a wolf. We have again provided you with the framework for this class, but you will have to implement the `step()` method. This method should move the Sheep using the `random_move()` method implemented earlier, then conditionally reproduce. You can use the 'new_agent()' method implemented in model to create new sheep.

Attributes/methods you might need:
 - `self.random_move()`
 - `self.model.new_agent()`
 - `self.model.sheep_reproduction_chance`

<!-- grade_id: cell-bdcef78c12f1fbe2, points: 0, grade: false, solution: true -->
```python
class Sheep(RandomWalker):
    def __init__(self, unique_id, model, pos):
        super().__init__(unique_id, model, pos)

    def step(self):
        '''
        This method should move the Sheep using the `random_move()` method implemented earlier, then conditionally reproduce.
        '''
        self.random_move()
        if random.random() < self.model.sheep_reproduction_chance:
            self.model.new_agent(Sheep, self.pos)

```

<!-- grade_id: cell-a38971467ab6ea79, points: 1, grade: true, solution: false -->
```python
tester = TestModel(10, 10)
tester.sheep_reproduction_chance = 1

# Create a Sheep 
start_position = (5, 5)
tester.new_agent(Sheep, start_position)

# Create a reference, so that we can properly test
sheep = tester.agents[0]

# Check movement
tester.step()
assert sheep.pos != (5, 5), "The sheep has not moved!"
assert abs(sheep.pos[0] - 5) <= 1 and abs(sheep.pos[1] - 5) <= 1, "The sheep has moved more than 1 square!"

# Check fertility
assert tester.n_agents == 2, "Your sheep is infertile!"
tester.sheep_reproduction_chance = 0
tester.step()
assert tester.n_agents == 2, "Your sheep are too fertile!"

# Check if the random_move function was used
old_pos = sheep.pos
sheep.random_move = lambda: False
tester.step()
assert old_pos == sheep.pos, "You didn't use the random_move() function!"
```

### Agents - Wolf
Each wolf has a fixed probability of dying at each timestep. When wolves meet a sheep they will eat it and then reproduce. We have again provided you with the framework for this class, but you will have to implement the step() method. This method should move the wolf, then check for sheep on its location, eat the sheep if it is there and reproduce, and finally conditionally die.


Attributes/methods you might need:
 - `self.random_move()`
 - `self.model.new_agent()`
 - `self.model.remove_agent()`
 - `self.model.wolf_death_chance`
 - [`MultiGrid.get_neighbors()`](https://github.com/projectmesa/mesa/blob/master/mesa/space.py)


<!-- grade_id: cell-e6b7ab29c84f80f5, points: 0, grade: false, solution: true -->
```python
class Wolf(RandomWalker):
    def __init__(self, unique_id, model, pos):
        super().__init__(unique_id, model, pos)

    def step(self):
        '''
        This method should move the wolf, then check for sheep on its location, 
        eat one sheep if it is available and reproduce, and finally, conditionally die.
        '''
        self.random_move()
        cellmates = self.model.grid.get_cell_list_contents([self.pos])
        sheep = [agent for agent in cellmates if isinstance(agent, Sheep)]
        if sheep:
            self.model.remove_agent(sheep[0])
            self.model.new_agent(Wolf, self.pos)
        if random.random() < self.model.wolf_death_chance:
            self.model.remove_agent(self)

```

<!-- grade_id: cell-7fa56d562b5c523c, points: 2, grade: true, solution: false -->
```python
tester = TestModel(10, 10)
tester.wolf_death_chance = 0

# Create a Wolf 
start_position = (5, 5)
tester.new_agent(Wolf, start_position)

# Create a reference, so that we can properly test
wolf = tester.agents[0]

# Check movement
tester.step()
assert wolf.pos != (5, 5), "The wolf has not moved!"
assert abs(wolf.pos[0] - 5) <= 1 and abs(wolf.pos[1] - 5) <= 1, "The wolf has moved more than 1 square!"

# Check death rates
assert tester.n_agents == 1, "Your wolf died!"
tester.wolf_death_chance = 1
tester.step()
assert tester.n_agents == 0, "Your wolf is invincible!"

# Create another model so that we can be sure the wolf eats the sheep
tester = TestModel(1, 1)
tester.wolf_death_chance = 0
tester.sheep_reproduction_chance = 0

# Create a Sheep and a Wolf
start_position = (0, 0)
tester.new_agent(Sheep, start_position)
tester.new_agent(Wolf, start_position)

# Check if the sheep was eaten
tester.step()
assert type(tester.agents[0]) is not Sheep, "The sheep has not been eaten!"
```

### Extension 1 - Predator Saturation (Crowding)

In the classic Lotka–Volterra model, predation is linear: 
if you double the number of predators, you double the predation rate.
In reality, predators interfere with each other when crowding.

**Task:** Extend your `Wolf.step()` method with a new WolfCrowding class below so that a wolf’s probability
of successfully eating a sheep depends on how many *wolves* are in the
same grid cell as the wolf.

Use `self.model.grid.get_cell_list_contents([self.pos])` to get all agents in the current cell, then:
1) Count how many of those agents are `Wolf` (including `self`)
2) Let `hunt_prob = 1.0 / n_wolves_in_cell`
3) Only eat a sheep if `random.random() < hunt_prob`

Everything else in `step()` should remain the same:
- Move first (still using `random_move()`)
- If the wolf successfully hunts, it eats one sheep and reproduces.
- Then it may die according to `self.model.wolf_death_chance`

```python
class WolfCrowding(RandomWalker):
    def __init__(self, unique_id, model, pos):
        super().__init__(unique_id, model, pos)

    def step(self):
        '''
        Wolf step method with predator saturation.
        '''
        self.random_move()
        cellmates = self.model.grid.get_cell_list_contents([self.pos])
        sheep = [agent for agent in cellmates if isinstance(agent, Sheep)]
        wolves = [agent for agent in cellmates if isinstance(agent, (Wolf, WolfCrowding))]
        hunt_prob = 1.0 / len(wolves)
        if sheep and random.random() < hunt_prob:
            self.model.remove_agent(sheep[0])
            self.model.new_agent(WolfCrowding, self.pos)
        if random.random() < self.model.wolf_death_chance:
            self.model.remove_agent(self)

```

```python
# Single wolf should always eat the sheep (hunt_prob = 1.0)
random.seed(0)
tester = TestModel(1, 1)
tester.wolf_death_chance = 0
tester.sheep_reproduction_chance = 0

# One sheep and one wolf in the same (only) cell
tester.new_agent(Sheep, (0, 0))
tester.new_agent(WolfCrowding, (0, 0))

tester.step()

sheep_remaining = any(isinstance(a, Sheep) for a in tester.agents)
assert not sheep_remaining, (
    "With exactly 1 wolf in the cell, the sheep should always be eaten "
    "(hunt_prob should be 1.0)."
)

# More wolves in the same cell -> lower hunting rate
random.seed(1)

def estimate_eat_rate(n_wolves, n_trials=3000):
    eaten = 0
    for _ in range(n_trials):
        m = TestModel(1, 1)
        m.wolf_death_chance = 0
        m.sheep_reproduction_chance = 0
        m.new_agent(Sheep, (0, 0))
        for _ in range(n_wolves):
            m.new_agent(WolfCrowding, (0, 0))
        before = sum(isinstance(a, Sheep) for a in m.agents)
        agents = m.agents
        rand_wolf = [a for a in agents if isinstance(a, WolfCrowding)][0]
        rand_wolf.step()
        after = sum(isinstance(a, Sheep) for a in m.agents)
        eaten += (before - after)
    return eaten / n_trials

rate_1 = estimate_eat_rate(1)
rate_4 = estimate_eat_rate(4)

# With hunt_prob = 1 / n_wolves, we expect roughly:
# rate_1 ≈ 1.0 and rate_4 ≈ 0.25, up to randomness.
assert rate_1 == 1.0, (
    f"Eat rate with 1 wolf is too low ({rate_1:.2f}); "
    "did you reduce hunt_prob even when only 1 wolf is present?"
)

assert round(rate_4, 2) == 0.25, (
    f"Eat rate with 4 wolves is too high ({rate_4:.2f}); "
    "crowding does not seem to reduce hunting success enough."
)

print(f"Crowding tests passed. Eat rate: 1 wolf={rate_1:.2f}, 4 wolves={rate_4:.2f}")
```

### Extension 2 - Movement Bias

So far, wolves move randomly (`random_move()`), which is inefficient.

**Task:** Implement a new class `TacticalWolf` that moves *towards*
nearby sheep instead of wandering randomly.

Rules for `TacticalWolf.step()`:

1. Look at all **neighboring cells** (Moore neighborhood, no center)
2. Among these neighbors, find all agents that are `Sheep` and collect the positions of the cells containing sheep
3. If there is at least one neighboring cell with a sheep, move into one of those cells (randomly if there are several options)
4. If no neighboring cell contains a sheep, fall back to `self.random_move()`
5. After moving, as usual, if there is a sheep in the current cell, eat one and reproduce, then die probabilistically.

Later, you can try plugging `TacticalWolf` into the `WolfSheep` model below instead of `Wolf` and see how much more efficient the predators become!

```python
class TacticalWolf(RandomWalker):
    def __init__(self, unique_id, model, pos):
        super().__init__(unique_id, model, pos)

    def step(self):
        '''
        Step method for TacticalWolf.
        '''
        possible_steps = self.model.grid.get_neighborhood(
            self.pos,
            moore=True,
            include_center=False,
        )
        sheep_steps = []
        for pos in possible_steps:
            cellmates = self.model.grid.get_cell_list_contents([pos])
            sheep = [agent for agent in cellmates if isinstance(agent, Sheep)]
            if sheep:
                sheep_steps.append(pos)
        if sheep_steps:
            new_position = random.choice(sheep_steps)
            self.model.grid.move_agent(self, new_position)
        else:
            self.random_move()
        cellmates = self.model.grid.get_cell_list_contents([self.pos])
        sheep = [agent for agent in cellmates if isinstance(agent, Sheep)]
        if sheep:
            self.model.remove_agent(sheep[0])
            self.model.new_agent(TacticalWolf, self.pos)
        if random.random() < self.model.wolf_death_chance:
            self.model.remove_agent(self)

```

```python
random.seed(1)

# TacticalWolf moves toward an adjacent sheep
tester = TestModel(10, 10)
tester.wolf_death_chance = 0
tester.sheep_reproduction_chance = 0

wolf_start = (5, 5)
sheep_pos = (5, 6) # directly adjacent

tester.new_agent(Sheep, sheep_pos)
tester.new_agent(TacticalWolf, wolf_start)

wolf = tester.agents[-1]

wolf.step()

assert wolf.pos == sheep_pos, (
    f"TacticalWolf should have moved to {sheep_pos} where the sheep was, "
    f"but is at {wolf.pos}."
)

print("Test 1 passed: TacticalWolf moved toward adjacent sheep.")

# TacticalWolf eats the sheep after moving onto its cell
tester = TestModel(10, 10)
tester.wolf_death_chance = 0
tester.sheep_reproduction_chance = 0

tester.new_agent(Sheep, sheep_pos)
tester.new_agent(TacticalWolf, wolf_start)

wolf = tester.agents[-1]
wolf.step()

remaining_sheep = [a for a in tester.agents if isinstance(a, Sheep)]
assert len(remaining_sheep) == 0, (
    "TacticalWolf moved to the sheep's cell but did not eat the sheep."
)

print("Test 2 passed: TacticalWolf ate the sheep after moving.")


# TacticalWolf falls back to random movement when no sheep nearby
tester = TestModel(10, 10)
tester.wolf_death_chance = 0
tester.sheep_reproduction_chance = 0

wolf_start = (5, 5)
# Put a sheep far away (not in the Moore neighborhood)
tester.new_agent(Sheep, (0, 0))
tester.new_agent(TacticalWolf, wolf_start)
wolf = tester.agents[-1]

tester.step()

assert wolf.pos != wolf_start, (
    "TacticalWolf did not move at all — random fallback movement is missing."
)
assert abs(wolf.pos[0] - 5) <= 1 and abs(wolf.pos[1] - 5) <= 1, (
    f"TacticalWolf moved more than 1 cell ({wolf.pos}) — "
    "movement should be constrained to neighboring cells."
)

print("Test 3 passed: TacticalWolf used random movement when no nearby sheep.")
```

### The model
Now that you have implemented the agents, we will introduce you to the model. The model links the agents to the environment and provides methods that make interactions with the model easy. It uses the `DataCollector` class, which we will explain later, and the `RandomActivation` class. The `RandomActivation` class is a scheduler that iterates through each of the agents randomly and uses its `step()` method. We have seperated the Sheep scheduler from the Wolf scheduler to prevent what is called a ["race-condition"](https://en.wikipedia.org/wiki/Race_condition) wherein a Sheep gets eaten before it is its turn to `step()`. 

At the time of writing of this notebook, there are 4 different schedulers available in Mesa:
- `BaseScheduler`: This is a scheduler that activates agents one at a time, in the order they were added.
- `RandomActivation`: Activates each agent once per step, in random order, with the order reshuffled every step.
- `SimultaneousActivation`: Activates all agents simultaneously.
- `StagedActivation`: Allows agent activation to be divided into several stages instead of a single `step` method.

Check out [the docs](https://mesa.readthedocs.io/en/stable/tutorials/intro_tutorial.html#adding-the-scheduler) to learn more.

```python
import random

from mesa import Model
from mesa.space import MultiGrid
from mesa.datacollection import DataCollector
from mesa.time import RandomActivation

class WolfSheep(Model):
    '''
    Wolf-Sheep Predation Model
    '''
    
    def __init__(self, height=20, width=20,
                 initial_sheep=100, initial_wolves=30,
                 sheep_reproduction_chance=0.05, wolf_death_chance=0.05):

        super().__init__()

        self.height = height
        self.width = width
        self.initial_sheep = initial_sheep
        self.initial_wolves = initial_wolves
        self.sheep_reproduction_chance = sheep_reproduction_chance
        self.wolf_death_chance = wolf_death_chance

        # Add a schedule for sheep and wolves seperately to prevent race-conditions
        self.schedule_Sheep = RandomActivation(self)
        self.schedule_Wolf = RandomActivation(self)

        self.grid = MultiGrid(self.width, self.height, torus=True)
        self.datacollector = DataCollector(
             {"Sheep": lambda m: self.schedule_Sheep.get_agent_count(),
              "Wolves": lambda m: self.schedule_Wolf.get_agent_count()})

        # Create sheep and wolves
        self.init_population(Sheep, self.initial_sheep)
        self.init_population(Wolf, self.initial_wolves)

        # This is required for the datacollector to work
        self.running = True
        self.datacollector.collect(self)

    def init_population(self, agent_type, n):
        '''
        Method that provides an easy way of making a bunch of agents at once.
        '''
        for i in range(n):
            x = random.randrange(self.width)
            y = random.randrange(self.height)

            self.new_agent(agent_type, (x, y))

    def new_agent(self, agent_type, pos):
        '''
        Method that creates a new agent, and adds it to the correct scheduler.
        '''
        agent = agent_type(self.next_id(), self, pos)

        self.grid.place_agent(agent, pos)
        getattr(self, f'schedule_{agent_type.__name__}').add(agent)

    def remove_agent(self, agent):
        '''
        Method that removes an agent from the grid and the correct scheduler.
        '''
        self.grid.remove_agent(agent)
        getattr(self, f'schedule_{type(agent).__name__}').remove(agent)

    def step(self):
        '''
        Method that calls the step method for each of the sheep, and then for each of the wolves.
        '''
        self.schedule_Sheep.step()
        self.schedule_Wolf.step()

        # Save the statistics
        self.datacollector.collect(self)

    def run_model(self, step_count=200):
        '''
        Method that runs the model for a specific amount of steps.
        '''
        for i in range(step_count):
            self.step()

```

You might have noticed that we are using a so called `MultiGrid`. Again, [Mesa provides us with multiple different types of environments](https://mesa.readthedocs.io/en/latest/tutorials/intro_tutorial.html#adding-space) for the agents to exist in:
- `ContinuousSpace`: Each agent can have an arbitrary position.
- `SingleGrid`: Grid where each cell contains at most one object.
- `MultiGrid`: Grid where each cell can contain more than one object.
- `HexGrid`: extends the grid to handle hexagonal neighbors.
- `NetworkGrid`: Network where each node contains zero or more agents.

We are using `MultiGrid` as we want multiple agents to be able to be in one cell.

Running the model is now as simple as:

```python
model = WolfSheep()
model.run_model()
```

### Data collector
So far, we have tested the model by checking variables manually. This is not very efficient, and requires us to make complicated loops if we need more than just the end results. Mesa provides a class which can handle the data collection and storage. This class was already added to the model above; `DataCollector`. From the Mesa docs:

"The data collector stores three categories of data: model-level variables, agent-level variables, and tables (which are a catch-all for everything else). Model- and agent-level variables are added to the data collector along with either the name of an attribute that the agents or models have, or a function for collecting some data. Model-level collection functions take a model object as an input, while agent-level collection functions take an agent object as an input. Both then return a value computed from the model or each agent at their current state. When the data collector’s collect method is called, with a model object as its argument, it applies each model-level collection function to the model, and stores the results in a dictionary, associating the current value with the current step of the model. Similarly, the method applies each agent-level collection function to each agent currently in the schedule, associating the resulting value with the step of the model, and the agent’s unique_id."

Simply said, the `DataCollector` adds data to its dictionary when you call `collect()` method. We have added two functions that keep track of the amount of Sheep and Wolf to the `DataCollector` and call collect after each `step()` that the model makes.

Run the following piece of code to see what happens.

```python
%matplotlib inline
    
model = WolfSheep()
model.run_model()

data = model.datacollector.get_model_vars_dataframe()
data.plot()
```

See if you can make a similar plot which displays the average number of sheep a wolf has eaten over its lifetime for each `step()` in the model. You will need to change the `step()` method of the wolf so that it counts every sheep it has eaten, and you will need to change the `DataCollector` to log this data. You can then use `model.datacollector.get_agent_vars_dataframe()` to get your data.

Take a look at the [Mesa Read The Docs](https://mesa.readthedocs.io/en/latest/tutorials/intro_tutorial.html#collecting-data) for more details.

```python
if not hasattr(Wolf, '_base_init'):
    Wolf._base_init = Wolf.__init__

def tracked_wolf_init(self, unique_id, model, pos):
    Wolf._base_init(self, unique_id, model, pos)
    self.sheep_eaten = 0

def tracked_wolf_step(self):
    self.random_move()
    cellmates = self.model.grid.get_cell_list_contents([self.pos])
    sheep = [agent for agent in cellmates if isinstance(agent, Sheep)]
    if sheep:
        self.model.remove_agent(sheep[0])
        self.sheep_eaten += 1
        self.model.new_agent(Wolf, self.pos)
    if random.random() < self.model.wolf_death_chance:
        self.model.remove_agent(self)

class CombinedSchedule:
    def __init__(self, model):
        self.model = model

    @property
    def steps(self):
        return self.model.schedule_Wolf.steps

    @property
    def agents(self):
        return list(self.model.schedule_Sheep.agents) + list(self.model.schedule_Wolf.agents)

Wolf.__init__ = tracked_wolf_init
Wolf.step = tracked_wolf_step

model = WolfSheep()
model.schedule = CombinedSchedule(model)
model.datacollector = DataCollector(
    {"Sheep": lambda m: m.schedule_Sheep.get_agent_count(),
     "Wolves": lambda m: m.schedule_Wolf.get_agent_count()},
    agent_reporters={"Breed": lambda a: type(a).__name__,
                     "Sheep eaten": lambda a: getattr(a, "sheep_eaten", 0)}
)
model.datacollector.collect(model)
model.run_model()

agent_data = model.datacollector.get_agent_vars_dataframe().reset_index()
wolf_data = agent_data[agent_data['Breed'] == 'Wolf']
wolf_data.groupby('Step')['Sheep eaten'].mean().plot(ylabel='average sheep eaten');

```

### Visualisation
Mesa provides you with an interactive visualisation toolkit. This is done through a local server that uses JavaScript to draw  everything in a browser window. An advantage of this is that we can run models step by step, spotting patterns, behaviors, bugs, developing new intuitions, hypotheses or insights. It is also a very good way to explain a model to your roommates or grandparents.

We have provided you with a seperate `server.py` that shows you how you could visualise this model, as the visualisation will not cooperate with notebooks. Feel free to take a look and change elements to get accustomed to how the visualisation module works. For more information, visit [this link](https://mesa.readthedocs.io/en/latest/tutorials/visualization_tutorial.html#building-your-own-visualization-component).

Run this model by issuing the following command in the command line (in your 1-mesa folder):

`> ipython server.py`

NOTE: Keep in mind that the "Frames Per Second" setting does not take into account the time it takes to build a frame. This means that when your simulation takes more time to build a frame than it takes the visualisation to display it, and you press stop, the visualisation will continue untill it has displayed all the frames that it should have displayed in the first place. To prevent this from happening, do not set "Frames Per Second" too high.

### Optional: Try `TacticalWolf` in the `WolfSheep` model

As an extra experiment, try modifying the `WolfSheep` model so it uses the `TacticalWolf` instead of `Wolf` for predators.

- Compare the population time series before and after.
- How much do you need to increase `sheep_reproduction_chance` to avoid prey extinction when using `TacticalWolf`?

### Mesa best practices
To ensure that our models are easily accessible and maintainable, Mesa provides us with some general principles. These principles can be found [here](https://mesa.readthedocs.io/en/latest/best-practices.html), but the main take-away is that you should structure your files. The model class for example, should be located in `model.py` and generally, the agents should be located in `agent.py`.
